# I AM NOT DONE

def practice():
    # TODO: Create a list named `fruits` containing "apple", "banana", "cherry"
    
    # TODO: Access the first item in `fruits` and assign it to `first_fruit`
    
    # TODO: Access the last item in `fruits` and assign it to `last_fruit`
    
    return locals()
