# I AM NOT DONE

def practice():
    # TODO: Create a variable `sum` equal to 5 plus 10
    
    # TODO: Create a variable `difference` equal to 100 minus 5
    
    # TODO: Create a variable `product` equal to 6 times 7
    
    # TODO: Create a variable `quotient` equal to 80 divided by 8
    
    return locals()
