# I AM NOT DONE

class Animal:
    def speak(self):
        return "Some sound"

# TODO: Define a class named `Cat` that inherits from `Animal`
    # TODO: Override the `speak` method to return "Meow"

def practice():
    # TODO: Create an instance of `Cat` and assign it to `my_cat`
    
    return locals()
