# I AM NOT DONE

# TODO: Define a class named `Dog`
    # TODO: Inside the class, define an `__init__` method that takes `self` and `name`
        # TODO: Assign the `name` argument to `self.name`
    
    # TODO: Define a method named `bark` that returns the string "Woof!"

def practice():
    # TODO: Create an instance of `Dog` with the name "Fido" and assign it to `my_dog`
    
    return locals()
