# I AM NOT DONE

def practice():
    # TODO: Create a dictionary named `person` with keys "name" (value "Alice") and "age" (value 30)
    
    # TODO: Access the value associated with the key "name" and assign it to `name`
    
    # TODO: Access the value associated with the key "age" and assign it to `age`
    
    return locals()
